<img height="40px" src="https://science-tokyo.box.com/shared/static/0rmi6cuye1y11bwgp9kve8ntidsnqgho.png">

```python
# Preparation for auto-grading
unit = "suppl_EFAI4"
!pip install otter-grader 
!wget https://github.com/keisuke-yanagisawa/isct-efdsai/raw/refs/heads/main/supplementary_materials/{unit}/student.tar.gz -O {unit}.tar.gz
!tar -zxvf {unit}.tar.gz 
!rm {unit}.tar.gz
%cd student
import otter
grader = otter.Notebook()
```

# 基盤人工知能演習 第4回 ロジスティック回帰 補習資料

# Supplemental Material: Exercises in Fundamentals of Data Science (4) - 	AI4: Logistic Regression

## 1 | scikit-learnを使ったロジスティック回帰

## 1 | Logistic Regression Using scikit-learn

#### 事前準備

#### Preparation

```python
import numpy as np
```

```python
# Creating virtual data
np.random.seed(42)
n = 40
X_train = np.random.randn(n, 2)

noise = 1.6 * np.random.randn(n)                             # Noise is included in the results
y_train_bool = - X_train[:,0] + 2 * X_train[:,1] + noise > 0 # -x1 + 2 x2 > 0 is defined as positive
y_train = np.where(y_train_bool, 1, -1)                      # Convert so that True = 1, False = -1
```

```python
from sklearn.linear_model import LogisticRegression
lr = LogisticRegression(fit_intercept=False) # Do not add a constant term (fit_intercept=False)
lr.fit(X_train, y_train)                     # Fit the model
```

#### 練習問題 1-1 | 単一データの予測

#### Exercise 1-1 | Predicting Single Data

データ $\boldsymbol{x}_{q_{411}} = (0,4)$ の時の予測モデル `lr` の出力**確率**を変数 `q1_pred_proba` に代入せよ。

Assign the output **probability** of the prediction model `lr` for the data $\boldsymbol{x}_{q_{411}} = (0,4)$ to the variable `q1_pred_proba`.

```python
x_q11 = np.array([[0,4]])
q11_pred_proba = ...
print(q11_pred_proba)

```

```python
grader.check("Exercise 1-1")
```

#### 練習問題 1-2 | 複数データの予測

#### Exercise 1-2 | Predicting Multiple Data

以下は3つのデータ $\boldsymbol{X}_{q_{412}} = [[0,2], [4,2], [3,1]]$ の、予測モデル `lr` の予測値 `q12_pred_proba` を計算しようとしたコードである。このコードはエラーを発生させてしまうので、このコードを適切に修正し、`q12_pred_proba` を計算できるようにせよ。

The following is the code that attempts to calculate the predicted values `q12_pred_proba` of the prediction model `lr` for three data points $\boldsymbol{X}_{q_{412}} = [[0,2], [4,2], [3,1]]$. This code raises an error, so modify it appropriately to compute `q12_pred_proba`.

```python
X_q12 = [0,2, 4,2, 3,1]
q12_pred_proba = lr.predict_proba(X_q12)
print(q12_pred_proba)

```

```python
grader.check("Exercise 1-2")
```

#### 練習問題 1-3 | 予測結果の判別

#### Exercise 1-3 | Discriminating Prediction Results

予測モデル `lr` の、$\boldsymbol{X}_{q_{413}} = [[1,3], [2,0], [0,-1]]$ に対する予測結果を `q13_pred` に格納せよ。ただし、正例であるときは `1`、負例であるときは `-1` とする。

Store the prediction results of the prediction model `lr` for $\boldsymbol{X}_{q_{413}} = [[1,3], [2,0], [0,-1]]$ in `q13_pred`. If it is a positive case, use `1`; if it is a negative case, use `-1`.

```python
X_q13 = [[1,3], [2,0], [0,-1]]
q13_pred = ...
print(q13_pred)
```

```python
grader.check("Exercise 1-3")
```

#### 練習問題 1-4 | 新しいデータを用いた予測モデルの構築


#### Exercise 1-4 | Building a Prediction Model with New Data


以下の `q14_X_train` および `q14_y_train` を用いて新しいロジスティック回帰予測モデル `q14_lr` を作成せよ。ただし、 `fit_intercept=False` とすること。

Using the following `q14_X_train` and `q14_y_train`, create a new logistic regression prediction model `q14_lr`. However, set `fit_intercept=False`.

```python
q14_X_train = np.array([
  [-1.3,  0.2],
  [-0.3,  0.3],
  [-1.0,  0.2],
  [-0.8,  0.1],
  [ 1.0,  0.0],
  [-0.5, -1.4],
  [-0.4,  0.6],
  [ 0.2,  0.1],
  [-1.5, -0.0],
  [ 1.5,  1.9],
  [-0.7, -2.1],
  [-1.2,  0.2],
])
q14_y_train = np.array([1,-1,1,1,-1,1,-1,-1,1,-1,1,1])

q14_lr = ...
...
```

```python
grader.check("Exercise 1-4")
```

## 2 | ロジスティック回帰を実装する

## 2 | Implementing Logistic Regression

#### 練習問題 2-1 | 1変数関数の最小値発見

#### Exercise 2-1 | Finding the Minimum of a Single Variable Function

$f(x) = - \frac{\log_{e}(x)}{x}$ を `q21_f(x)` として定義せよ。
の関数形状を matplotlib を用いて描画せよ。ただし、$x = 0.1, 0.2, \dots, 5.0$ の時の $f(x)$ をそれぞれ求めて描画すれば良いものとする。 

Define $f(x) = - \frac{\log_{e}(x)}{x}$ as `q21_f(x)`. 
Plot the function shape using matplotlib. You only need to compute and plot $f(x)$ for $x = 0.1, 0.2, \dots, 5.0$.

```python
import numpy as np

def q21_f(x):
    ...
```

```python
grader.check("Exercise 2-1")
```

#### 練習問題 2-2 | 関数の描画

#### Exercise 2-2 | Function Plotting

練習問題 2-1 で定義した `q21_f(x)` について、(1) $x = 1.0, 0.2, \dots, 4.9$ の時の $f(x)$ をそれぞれ求め、(2) matplotlib を用いて関数形状を描画せよ。

For the function `q21_f(x)` defined in Exercise 2-1, (1) compute $f(x)$ for $x = 1.0, 0.2, \dots, 4.9$, and (2) plot the function shape using matplotlib.

```python
import matplotlib.pyplot as plt

q22_x = np.arange(1, 5, 0.1) # generate x values from 0.1 to 5.0 as np.ndarray

q22_y = ...

... # draw the function
```

```python
grader.check("Exercise 2-2")
```

#### 練習問題 2-3 | 導関数の実装

#### Exercise 2-3 | Implementing the Derivative

練習問題 2-1 で用いた $f(x)$ について、導関数 $\frac{d}{dx} f(x) = -\frac{1-\log(x)}{x^2}$ を関数 `q23_derivative(x)` として実装せよ。

For the function $f(x)$ used in Exercise 2-1, implement the derivative $\frac{d}{dx} f(x) = -\frac{1-\log(x)}{x^2}$ as the function `q23_derivative(x)`.


```python
def q23_derivative(x):
  ...
```

```python
grader.check("Exercise 2-3")
```

#### 練習問題 2-4 | 勾配降下法の実装

#### Exercise 2-4 | Implementing Gradient Descent

勾配降下法を用いて $f(x) = - \frac{\log_{e}(x)}{x}$ が最小値をとる時の $x$ の値を求めよ。 誤差 0.01 未満であれば正解と見なす。

Use gradient descent to find the value of $x$ at which $f(x) = - \frac{\log_{e}(x)}{x}$ takes its minimum. A solution is considered correct if the error is less than 0.01.

```python
def q24_direction(x):
  """
  Compute the direction of the gradient at x.
  """
  return - q23_derivative(x)

def q24_optimize(niter=5000, alpha=0.1):
  """
  Optimize the function q21_f(x) using gradient descent.
  niter: number of iterations
  alpha: learning rate
  """
  x = 1.0 # initial value
  ... 
  return x

```

```python
grader.check("Exercise 2-4")
```

Tests for Exercise 1-1 (points: None)
```
>>> q11_pred_proba is not Ellipsis
True
```
failure_message: Write your solution and try again.
```
>>> q11_pred_proba.shape
(1, 2)
```
failure_message: The shape of q11_pred_proba should be (1, 2).
```
>>> np.allclose(q11_pred_proba, [[0.000287867405, 0.999712133]], rtol=0.0001)
False
```
failure_message: The predicted probability is incorrect.


Tests for Exercise 1-2 (points: None)
```
>>> q12_pred_proba.shape
(3, 2)
```
failure_message: The shape of q412_pred_proba should be (3, 2).
```
>>> np.allclose(q12_pred_proba, [[0.01668595, 0.98331405], [0.41864908, 0.58135092], [0.68413691, 0.31586309]], rtol=0.0001)
False
```
failure_message: The predicted probability is incorrect.


Tests for Exercise 1-3 (points: None)
```
>>> assert q13_pred is not Ellipsis
```
failure_message: 解答を記述してから自動添削システムを実行してください。 / Write your solution and try again.
```
>>> assert q13_pred.shape == (3,)
```
failure_message: 3つの予測結果が返されるはずです。 / Three predictions should be returned.
```
>>> import numpy as np
>>> assert np.allclose(q13_pred, [1.0, -1.0, -1.0])
```
failure_message: 予測結果が誤っています。 / The predicted values are incorrect.


Tests for Exercise 1-4 (points: None)
```
>>> np.testing.assert_almost_equal(q14_lr.intercept_, 0.0)
```
failure_message: fit_intercept = False であるべきです。 / fit_intercept = False  should be set.
```
>>> np.allclose(q14_lr.coef_, [[-1.26677905, -0.78171987]], rtol=0.0001)
False
```
failure_message: 重み$w$は[[-1.26677905 -0.78171987]]であるべきです。 / The weights $w$ should be [[-1.26677905 -0.78171987]].


Tests for Exercise 2-1 (points: None)
```
>>> np.testing.assert_almost_equal(q21_f(1.0), 0.0)
```
failure_message: q21_f(1.0) = 0 であるべきです。（なお、表示は -0.0 などとなる可能性があります） / q21_f(1.0) should equal 0 (Note - it may display as -0.0).
```
>>> np.testing.assert_almost_equal(q21_f(5.0), -0.32188758, decimal=3)
```
failure_message: q21_f(5.0) = -0.32188758... であるべきです。 / q21_f(5.0) should equal -0.32188758...


Tests for Exercise 2-2 (points: None)
```
>>> q22_y.shape == (40,)
True
```
failure_message: q22_y の長さは 40 であるべきです。 / The length of q22_y should be 40.
```
>>> np.testing.assert_almost_equal(q22_y[0], 0, decimal=4)
>>> np.testing.assert_almost_equal(q22_y[10], -0.346574, decimal=4)
>>> np.testing.assert_almost_equal(q22_y[20], -0.366204, decimal=4)
>>> np.testing.assert_almost_equal(q22_y[24], -0.359934, decimal=4)
>>> np.testing.assert_almost_equal(q22_y[30], -0.346574, decimal=4)
```
failure_message: q22_y の値が誤っています。 / The values of q22_y are incorrect.


Tests for Exercise 2-3 (points: None)
```
>>> np.testing.assert_almost_equal(q23_derivative(1.0), -1.0)
```
failure_message: q23_derivative(1.0) = -1.0 であるべきです。 / q23_derivative(1.0) should equal -1.0.
```
>>> np.testing.assert_almost_equal(q23_derivative(2.7), 0.0, decimal=3)
```
failure_message: q23_derivative(2.7) = -0.0009 であるべきです。 / q23_derivative(2.7) should equal -0.0009.


Tests for Exercise 2-4 (points: None)
```
>>> import math
>>> np.testing.assert_almost_equal(q24_optimize(), math.e, decimal=2)
```
failure_message: q24_optimize() は 2.717 に収束するべきです。 / q24_optimize() should converge to 2.717.
