<img height="40px" src="https://science-tokyo.box.com/shared/static/mst450uunryw194o872jt80qqk27lqux.png">

```python
# Preparation for auto-grading
!pip install otter-grader 
import otter
grader = otter.Notebook()
```

# 基盤人工知能演習 第2回 NumPyによる線形代数 補習資料

# Supplemental Material: Exercises in Fundamentals of Artificial Intelligence (2) - Introduction of NumPy

## 1 | NumPy配列

## 1 | NumPy Arrays

#### 練習問題 1-1 | 単位ベクトルの作成

#### Exercise 1-1 | Create Unit Vector

与えられた 3次元ベクトル `v = (x,y,z)` の単位ベクトルを求める関数 `normalize_vector(v)` を作成せよ。ただし、`v` は `np.array((x,y,z))` で作成されたNumPy配列であり、零ベクトルではないことを仮定して良い。関数の出力はNumPy配列であるようにせよ。

Create a function `normalize_vector(v)` that computes the unit vector of a given 3-dimensional vector `v = (x,y,z)`. Assume that `v` is a NumPy array created with `np.array((x,y,z))` and that it is not the zero vector. The function's output should also be a NumPy array.

```python
import numpy as np
import numpy.typing as npt

def normalize_vector(v: npt.NDArray[np.float64]) -> npt.NDArray[np.float64]:
  ... 
```

```python
grader.check("Exercise 1-1")
```

#### 練習問題 1-2 | 列方向・行方向の集計(1)

#### Exercise 1-2 | Column-wise and Row-wise Aggregation (1)

与えられた行列 `mat` の各行の最大値を取得して配列として返す関数 `get_max_per_row(mat)` を実装せよ。 ただし、入出力はいずれも NumPy 配列で表現されているものとする。

Implement a function `get_max_per_row(mat)` that retrieves the maximum value from each row of a given matrix `mat` and returns it as an array. Assume that both input and output are represented as NumPy arrays.

```python
import numpy as np
import numpy.typing as npt

def get_max_per_row(mat: npt.NDArray[np.float64]) -> npt.NDArray[np.float64]:
  ... 
```

```python
grader.check("Exercise 1-2")
```

#### 練習問題 1-3 | 列方向・行方向の集計(2)

#### Exercise 1-3 | Column-wise and Row-wise Aggregation (2)

与えられた行列 `mat` の各列の最小値を取得して配列として返す関数 `get_min_per_column(mat)` を実装せよ。 ただし、入出力はいずれも NumPy 配列で表現されているものとする。

Implement a function `get_min_per_column(mat)` that retrieves the minimum value from each column of a given matrix `mat` and returns it as an array. Assume that both input and output are represented as NumPy arrays.

```python
import numpy as np
import numpy.typing as npt

def get_min_per_column(mat: npt.NDArray[np.float64]) -> npt.NDArray[np.float64]:
  ... 
```

```python
grader.check("Exercise 1-3")
```

#### 練習問題 1-4 | 行列積演算の練習


#### Exercise 1-4 | Practice with Matrix Multiplication


2行4列の行列 `q14_X` と4行3列の行列 `q14_Y` を自由に準備し、その行列積 `q14_XY` を求めよ。

Prepare a 2×4 matrix `q14_X` and a 4×3 matrix `q14_Y`, and compute their matrix product `q14_XY`.

```python
q14_X = ...
q14_Y = ...
q14_XY = ...
```

```python
grader.check("Exercise 1-4")
```

#### 練習問題 1-5 | ベクトルの L2-Norm の算出

#### Exercise 1-5 | Calculation of the L2 Norm of a Vector

与えられた $N$ 次元ベクトル $v$ の L2ノルム $||v||_2 = \sqrt{x_1^2 + \cdots + x_N^2}$ を計算する関数 `calc_l2norm(v)` を、`dot()` を用いて実装せよ。（ `np.linalg.norm(v)` で計算可能であるが、これは利用しないこと）

Implement a function `calc_l2norm(v)` that calculates the L2 norm of a given N-dimensional vector $v$ as $||v||_2 = \sqrt{x_1^2 + \cdots + x_N^2}$ using `dot()`. (While this can be computed using `np.linalg.norm(v)`, do not use it.)

```python
import numpy as np

def calc_l2norm(v: npt.NDArray[np.float64]) -> np.float64:
  ... 
```

```python
grader.check("Exercise 1-5")
```

#### 練習問題 1-6 | 逆行列の逆行列

#### Exercise 1-6 | Inverse of the Inverse Matrix

行列 `q16_A` の逆行列の逆行列 `inv_inv_q16_A` を求めよ。また、これら2つがほぼ一致することを確認せよ。

Calculate the inverse of the inverse matrix `inv_inv_q16_A` of the matrix `q16_A`. Also, confirm that these two are approximately equal.

```python
q16_A = ...
inv_inv_q16_A = ...
```

```python
grader.check("Exercise 1-6")
```

#### 練習問題 1-7 | 対角行列の固有値・固有ベクトル

#### Exercise 1-7 | Eigenvalues and Eigenvectors of Diagonal Matrices

対角行列 `q17_diagX` の固有値 `q17_eig_vals` および固有ベクトル `q17_eig_vecs` を計算せよ。ただし、 `q17_eig_vals` は複数の固有値をまとめたNumPy配列、 `q17_eig_vecs` は各固有ベクトルを行とするNumPy配列による行列とすること。

Calculate the eigenvalues `q17_eig_vals` and eigenvectors `q17_eig_vecs` of the diagonal matrix `q17_diagX`. Assume that `q17_eig_vals` is a NumPy array containing multiple eigenvalues, and `q17_eig_vecs` is a matrix with each eigenvector as a row.

```python
q17_diagX = np.array([[1,0,0],[0,3,0],[0,0,-2]])
... 
print(q17_eig_vecs)
```

```python
grader.check("Exercise 1-7")
```

#### 練習問題 1-8 | for文での行列積計算と .dot() での行列積計算との速度差

#### Exercise 1-8 | Speed Difference Between Matrix Multiplication Using for Loops and .dot()

与えられた巨大な行列 `q18_mat1`, `q18_mat2` の行列積を計算するコードを以下に記述している。以下のコードを実行し、その速度差を確認せよ。

なお、この練習問題は穴埋め問題ではないため、自動添削システムは含まれていない。

Write the code below to calculate the matrix product of the large matrices `q18_mat1` and `q18_mat2`. Run the following code to check the speed difference.

Note that this exercise is not a fill-in-the-blank question, so the automatic grading system is not included.

```python
import time

size = 200 # the size of the matrices

q18_mat1 = np.random.rand(size, size)
q18_mat2 = np.random.rand(size, size)

# with Python loop
start = time.time()
q18_forloop = np.zeros((size, size))
for i in range(size):
  for j in range(size):
    for k in range(size):
      q18_forloop[i, j] += q18_mat1[i, k] * q18_mat2[k, j]
end = time.time()
print("For loop: ", end - start, "sec.")

# with NumPy dot
start = time.time()
q18_dot = q18_mat1.dot(q18_mat2)
end = time.time()
print("Dot     : ", end - start, "sec.")

```

Tests for Exercise 1-1 (points: None)
```
>>> sample_vector = np.array([1, 2])
>>> isinstance(normalize_vector(sample_vector), np.ndarray)
True
```
failure_message: 返り値がNumPy配列になっていません。 / The return value is not a NumPy array.
```
>>> sample_normed_vector = np.array([1, 0])
>>> normalize_vector(sample_normed_vector) == sample_normed_vector
array([ True,  True])
```
failure_message: 出力されたベクトルが単位ベクトルになっていません。 / The output vector is not a unit vector.
```
>>> sample_vector = np.array([3, 4])
>>> np.testing.assert_almost_equal(np.linalg.norm(normalize_vector(sample_vector)), 1.0)
```
failure_message: 出力されたベクトルが単位ベクトルになっていません。 / The output vector is not a unit vector.


Tests for Exercise 1-2 (points: None)
```
>>> sample_matrix = np.array([[1, 2, 3], [4, 5, 6]])
>>> isinstance(get_max_per_row(sample_matrix), np.ndarray)
True
```
failure_message: 返り値がNumPy配列になっていません。 / The return value is not a NumPy array.
```
>>> sample_matrix = np.array([[1, 2], [3, 4], [5, 6]])
>>> len(get_max_per_row(sample_matrix)) == 3
True
```
failure_message: 出力されたベクトルの次元が3ではありません / The output vector's dimension is not 3
```
>>> sample_matrix = np.array([[1]])
>>> get_max_per_row(sample_matrix) == 1
array([ True])
```
failure_message: 出力されたベクトルが各行の最大値を保持していません / The output vector does not contain the maximum value of each row
```
>>> sample_matrix = np.array([[1, 2], [3, 4], [6, 5]])
>>> get_max_per_row(sample_matrix) == [2, 4, 6]
array([ True,  True,  True])
```
failure_message: 出力されたベクトルが各行の最大値を保持していません / The output vector does not contain the maximum value of each row


Tests for Exercise 1-3 (points: None)
```
>>> sample_matrix = np.array([[1, 2, 3], [4, 5, 6]])
>>> isinstance(get_min_per_column(sample_matrix), np.ndarray)
True
```
failure_message: 返り値がNumPy配列になっていません。 / The return value is not a NumPy array.
```
>>> sample_matrix = np.array([[1, 2], [3, 4], [5, 6]])
>>> len(get_min_per_column(sample_matrix)) == 2
True
```
failure_message: 出力が入力の列の次元と一致していません / The output's dimension does not match the input's column dimension
```
>>> sample_matrix = np.array([[1]])
>>> get_min_per_column(sample_matrix) == sample_matrix
array([[ True]])
```
failure_message: 出力されたベクトルが各列の最小値を保持していません / The output vector does not contain the minimum value of each column
```
>>> sample_matrix = np.array([[1, 3], [4, 2], [5, 6]])
>>> get_min_per_column(sample_matrix) == [1, 2]
array([ True,  True])
```
failure_message: 出力されたベクトルが各列の最小値を保持していません / The output vector does not contain the minimum value of each column


Tests for Exercise 1-4 (points: None)
```
>>> q14_X.shape
(2, 4)
```
failure_message: 次元が異なっています / The dimensions are different.
```
>>> q14_Y.shape
(4, 3)
```
failure_message: 次元が異なっています / The dimensions are different.
```
>>> q14_XY.shape
(2, 3)
```
failure_message: 次元が異なっています / The dimensions are different.


Tests for Exercise 1-5 (points: None)
```
>>> sample_vector = np.array([1.0, 2.0, 3.0])
>>> np.isscalar(calc_l2norm(sample_vector))
True
```
failure_message: 出力がスカラーではありません / The output is not a scalar
```
>>> sample_vector = np.array([1.0, 2.0])
>>> np.testing.assert_almost_equal(calc_l2norm(sample_vector), np.linalg.norm(sample_vector))
```
failure_message: 2次元ベクトルに対するL2ノルムの計算が正しくありません / The calculation of the L2 norm for a 2-dimensional vector is incorrect
```
>>> sample_vector = np.array([1.0, 2.0, 3.0])
>>> np.testing.assert_almost_equal(calc_l2norm(sample_vector), np.linalg.norm(sample_vector))
```
failure_message: 3次元ベクトルに対するL2ノルムの計算が正しくありません / The calculation of the L2 norm for a 3-dimensional vector is incorrect
```
>>> sample_vector = np.array([0, 0, 0, 0])
>>> np.testing.assert_almost_equal(calc_l2norm(sample_vector), np.linalg.norm(sample_vector))
```
failure_message: 零ベクトルに対するL2ノルムの計算が正しくありません / The calculation of the L2 norm for a zero vector is incorrect


Tests for Exercise 1-6 (points: None)
```
>>> inv_inv_q16_A.shape == q16_A.shape
True
```
failure_message: 3*3次元である必要があります / The dimension should be 3*3
```
>>> np.allclose(inv_inv_q16_A, q16_A)
True
```
failure_message: 逆行列の逆行列が元の行列と一致しません。 / The inverse of the inverse matrix does not match the original matrix.


Tests for Exercise 1-7 (points: None)
```
>>> q17_eig_vals.shape
(3,)
```
failure_message: 次元数は3である必要があります / The dimension should be 3
```
>>> q17_eig_vecs.shape
(3, 3)
```
failure_message: 次元数は3*3である必要があります / The dimension should be 3*3
```
>>> assert sorted(q17_eig_vals) == [-2.0, 1.0, 3.0]
```
failure_message: 固有値は -2, 1, 3 の3つになっている必要があります / The eigenvalues must be -2, 1, 3
```
>>> for _, (val, vec) in enumerate(zip(q17_eig_vals, q17_eig_vecs.T)):
...     np.testing.assert_allclose(q17_diagX @ vec, val * vec)
```
failure_message: 各固有値と固有ベクトルが対応していません / The eigenvalues do not correspond to the eigenvectors
