<img height="40px" src="https://science-tokyo.box.com/shared/static/jneg12ikt69n0j5q77qidm7g3a1iwa9b.png">

```python
# Preparation for auto-grading
!pip install otter-grader 
import otter
grader = otter.Notebook()
```

# 基盤データサイエンス演習 第5回 主成分分析 (principal component analysis: PCA) 補習資料

# Supplemental Material: Excercises in Fundamentals of Data Science (5) - Principal Component Analysis (PCA)

## 1 | データセットの確認

## 1 | Checking the Dataset

#### 事前準備

#### Preliminaries

```python
from sklearn import datasets
import numpy as np

d = datasets.load_digits() # load "digits" dataset
```

#### 練習問題 1-1 | データ数の確認

#### Exercise 1-1 | Check the Number of Data Points

`d.target` には、それぞれのデータがどのクラスに属しているかが記録されている。クラス0に属するデータの数を求め、`count_labelzero` に記録せよ。

`d.target` records which class each data point belongs to. Count the number of data points belonging to class 0 and record it in `count_labelzero`.

```python
count_labelzero = 0
for label in d.target:
    ...
"""
print(count_labelzero)
```

```python
grader.check("Exercise 1-1")
```

#### 練習問題 1-2 | 1つのデータが持つ特徴量数

#### Exercise 1-2 | Number of Features for One Data Point

与えられたデータセット（変数の型は行列形式のNumPy配列 `np.ndarray` であることを仮定してよい）内の各データが持つ特徴量数（次元数）は一定のはずである。特徴量数を求める関数 `get_n_features(X)` を作成せよ。

Assuming that the provided dataset is a matrix in the form of a NumPy array `np.ndarray`, the number of features (dimensions) for each data point should be constant. Create a function `get_n_features(X)` to determine the number of features.

```python
def get_n_features(X):
    ...
```

```python
grader.check("Exercise 1-2")
```

#### 練習問題 1-3 | ある特徴量の平均値の計算

#### Exercise 1-3 | Calculate the Mean of a Feature

digitsデータセット `d.data` の各データは64次元の特徴量を持ち、そのようなデータが1000件以上存在している。このうち、10番目の特徴量 `x[10]` について、全データにおける平均値を求めよ。

Each data point in the digits dataset `d.data` has 64 features, and there are more than 1000 such data points. Calculate the mean of the 10th feature `x[10]` across all data points.

```python
...

print(q13_x10_average)
```

```python
grader.check("Exercise 1-3")
```

## 2 | 主成分分析 (PCA)

## 2 | Principal Component Analysis (PCA)

#### 練習問題 2-1 | 主成分分析の実行

#### Exercise 2-1 | Perform Principal Component Analysis

digitsデータセット `d.data` に対して主成分分析を実行せよ。第1主成分の寄与率を求め、`q21_pc1_ratio` に記録せよ。ただし、寄与率は `q21_pca.components_` で求めることができる。


Perform principal component analysis on the digits dataset `d.data`. Calculate the contribution rate of the first principal component and record it in `q21_pc1_ratio`. The contribution rate can be calculated using `q21_pca.components_`.

```python
from sklearn.decomposition import PCA

q21_X = d.data
q21_pca = ...
...
q21_pc1_ratio = ...
print(q21_pc1_ratio)
```

```python
grader.check("Exercise 2-1")
```

#### 練習問題 2-2 | 各主成分同士の直交性

#### Exercise 2-2 | Orthogonality of Principal Components

PCA の各主成分同士は直交するはずである。`q21_pca.components_` に格納されている主成分ベクトル同士が直交しているかどうかを確認せよ。
2つのベクトルの内積が0であればそれらのベクトルは直交している、という事実を利用して、2つのベクトルが直交しているか否かを判定する関数 `q22_is_orthogonal(v1, v2)` を作成し、それを用いて確認すること。

ただし、実数値の計算誤差を考慮して、内積が0に等しいかどうかを判定する際には、`np.isclose(val, 0)` を用いること。

```otter

The principal components from PCA should be orthogonal to each other. Check whether the principal component vectors stored in `q21_pca.components_` are orthogonal.
Utilize the fact that if the dot product of two vectors is 0, the vectors are orthogonal. Create a function `q22_is_orthogonal(v1, v2)` to determine if two vectors are orthogonal, and use it to verify this.

However, to account for floating-point computation errors, use `np.isclose(val, 0)` when determining if the dot product is equal to 0.

```otter

```python
def q22_is_orthogonal(v1, v2):
    ...
    return np.isclose(..., 0)
```

```python
grader.check("Exercise 2-2")
```

#### 練習問題 2-3 | 主成分得点の計算

#### Exercise 2-3 | Calculate Principal Component Scores

各データの主成分得点は、各主成分と各データとの内積によって得られる。`q21_pca.components_` を用いて、データの42番目の要素`q21_X[42]` の第1主成分の主成分得点を計算し、 `q23_pc1_score` に記録せよ。

The principal component scores for each data point are obtained by taking the dot product between each principal component and each data point. Using `q21_pca.components_`, calculate the first principal component score for the 42nd data element `q21_X[42]` and record it in `q23_pc1_score`.

```python
q23_pc1_score = np.dot(..., ...)
```

```python
grader.check("Exercise 2-3")
```

Tests for Exercise 1-1 (points: None)
```
>>> count_labelzero == 178
True
```
failure_message: count_labelzero の値が正しくありません。 / The value of count_labelzero is incorrect.


Tests for Exercise 1-2 (points: None)
```
>>> assert isinstance(get_n_features(d.data), np.integer) or isinstance(get_n_features(d.data), int)
```
failure_message: get_n_features() の返り値が整数型ではありません。return 文は書きましたか？ / The return value of get_n_features() is not an integer. Did you write the return statement?
```
>>> assert get_n_features(np.random.rand(100, 20)) == 20
```
failure_message: 20種類の特徴量を持つデータ100件からなる X を入力したときの出力が 20 になっていません。 / The output for X consisting of 100 data points with 20 features is not 20.
```
>>> assert get_n_features(np.random.rand(83, 15)) == 15
```
failure_message: 15種類の特徴量を持つデータ83件からなる X を入力したときの出力が 15 になっていません。 / The output for X consisting of 83 data points with 15 features is not 15.
```
>>> assert get_n_features(np.random.rand(1, 20)) == 20
```
failure_message: 20種類の特徴量を持つデータ1件からなる X を入力したときの出力が 20 になっていません。 / The output for X consisting of 1 data point with 20 features is not 20.
```
>>> assert get_n_features(np.random.rand(20, 1)) == 1
```
failure_message: 1種類の特徴量を持つデータ20件からなる X を入力したときの出力が 1 になっていません。 / The output for X consisting of 20 data points with 1 feature is not 1.


Tests for Exercise 1-3 (points: None)
```
>>> assert np.isclose(q13_x10_average, 10.382303839732888, atol=1e-06)
```
failure_message: q13_x10_average の値が正しくありません。 / The value of q13_x10_average is incorrect.


Tests for Exercise 2-1 (points: None)
```
>>> from sklearn.utils.validation import check_is_fitted
>>> check_is_fitted(q21_pca)
```
failure_message: q21_pca.fit() が実行されていません。 / q21_pca.fit() has not been executed.
```
>>> assert np.isclose(q21_pc1_ratio, 0.14890593593061524, rtol=1e-05)
```
failure_message: q21_pc1_ratio の値が正しくありません。 / The value of q21_pc1_ratio is incorrect.


Tests for Exercise 2-2 (points: None)
```
>>> assert isinstance(q22_is_orthogonal(np.array([1, 0]), np.array([0, 1])), np.bool_) or isinstance(q22_is_orthogonal(np.array([1, 0]), np.array([0, 1])), bool)
```
failure_message: q22_is_orthogonal() の返り値が真偽値型ではありません。return 文は書きましたか？ / The return value of q22_is_orthogonal() is not a boolean. Did you write the return statement?
```
>>> assert q22_is_orthogonal(np.array([1, 0]), np.array([0, 1])) == True
```
failure_message: q22_is_orthogonal([1, 0], [0, 1]) の返り値が True になっていません。 / The return value of q22_is_orthogonal([1, 0], [0, 1]) is not True.
```
>>> assert q22_is_orthogonal(np.array([1, 0]), np.array([1, 0])) == False
```
failure_message: is_orthogonal([1, 0], [1, 0]) の返り値が False になっていません。 / The return value of q22_is_orthogonal([1, 0], [1, 0]) is not False.
```
>>> from sklearn.decomposition import PCA
>>> q22_pca = PCA(n_components=2).fit(q21_X)
>>> assert q22_is_orthogonal(q22_pca.components_[0], q22_pca.components_[1])
```
failure_message: pca.components_ が直交していないと判定されています。 / It is determined that pca.components_ are not orthogonal.


Tests for Exercise 2-3 (points: None)
```
>>> assert isinstance(q23_pc1_score, (int, float))
```
failure_message: 主成分得点はスカラー値であり、ベクトルではありません。 / The principal component score should be a scalar, not a vector.
```
>>> assert np.isclose(q23_pc1_score, np.dot(q21_X[42], q21_pca.components_[0]))
```
failure_message: 値が誤っています。内積は np.dot(v1, v2) で求めることができます。 / The value is incorrect. The dot product can be calculated with np.dot(v1, v2).
