<img height="40px" src="https://science-tokyo.box.com/shared/static/vfv8j6e0fju9huxspsgjpbt350f75zjj.png">

```python
# Preparation for auto-grading
!pip install otter-grader 
import otter
grader = otter.Notebook()
```

# 基盤人工知能演習 第3回 線形回帰と多項式回帰 補習資料

# Supplemental Material: Exercises in Fundamentals of Artificial Intelligence - AI3: Linear and Polynomial Regression

## 1 | 線形回帰・Ridge回帰

## 1 | Linear Regression and Ridge Regression

#### 練習問題 1-1 | 入力データセットの作成

#### Exercise 1-1 | Create the Input Dataset

$y = x_1 + x_2 + 1$ となるような10件のサンプルデータを自由に作成せよ。ただし、10件全てにおいて $y$ が異なる値になるようにせよ。
各データは `q11_X` および `q11_y` という変数に、NumPy配列として格納すること。

Create 10 samples such that $y = x_1 + x_2 + 1$, ensuring that each sample has a unique $y$ value. Store the data in variables `q11_X` and `q11_y` as NumPy arrays.

```python
import numpy as np

... 

print(q11_X)
print(q11_y)
```

```python
grader.check("Exercise 1-1")
```

#### 練習問題 1-2 | 線形回帰の実行

#### Exercise 1-2 | Perform Linear Regression

練習問題 1-1 で作成した `q11_X` および `q11_y` を使って線形回帰モデル `q12_lr` を構築し、その係数を確認せよ。

Using the `q11_X` and `q11_y` created in Exercise 1-1, build a linear regression model `q12_lr` and check its coefficients.

```python
from sklearn.linear_model import LinearRegression

q12_lr = ...
... 

print(q12_lr.coef_)
print(q12_lr.intercept_)
```

```python
grader.check("Exercise 1-2")
```

#### 練習問題 1-3 | Ridge回帰の実行

#### Exercise 1-3 | Perform Ridge Regression

`q11_X` および `q11_y` を使って、$\alpha=1$ としたときのRidge回帰モデル `q13_ridge_1` と、$\alpha=10$ としたときのRidge回帰モデル `q13_ridge_10` をそれぞれ構築し、その係数を比較せよ。

Using `q11_X` and `q11_y`, build Ridge regression models `q13_ridge_1` with $\alpha=1$ and `q13_ridge_10` with $\alpha=10$, and compare their coefficients.

```python
from sklearn.linear_model import Ridge

... 

print("alpha =  1", q13_ridge_1.coef_, q13_ridge_1.intercept_)
print("alpha = 10", q13_ridge_10.coef_, q13_ridge_10.intercept_)
```

```python
grader.check("Exercise 1-3")
```

## 2 | 多項式回帰とその応用

## 2 | Polynomial Regression and Its Applications

#### 事前準備

#### Preparation

今回用いるデータ `x` および `y` は、 $y = \cos(x)$ となるようなデータである（$x$は [rad] とし、ノイズは一切含まれていないことを仮定する）。この数式の関係であることを念頭に置きながら、この後の練習問題に取り組んで欲しい。

The data used here, `x` and `y`, follows $y = \cos(x)$, where $x$ is in [rad], and it is assumed that no noise is present. Keep this in mind as you tackle the following exercises.

```python
import numpy as np
from sklearn.model_selection import train_test_split

x = np.random.rand(80)*8
y = np.cos(x)

x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=0.5, random_state=0)
```

```python
# check the data

import matplotlib.pyplot as plt
plt.plot(np.linspace(0, 8, 100), np.cos(np.linspace(0, 8, 100)), label="$\cos(x)$", c="green")
plt.scatter(x_train, y_train, label="train", c="blue")
plt.scatter(x_test, y_test, label="test", c="red")
plt.xlabel("x [rad]")
plt.ylabel("y")
plt.title("$y = \cos(x)$")
plt.legend()
plt.show()
```

#### 練習問題 2-1 | `PolynomialFeatures` 相当の関数の作成

#### Exercise 2-1 | Create a Function Equivalent to `PolynomialFeatures`

`sklearn.preprocessing.polynomialFeatures` は、多項式回帰を行うための特徴量生成を行うものであった。
同様に、NumPy 1次元配列（例：`np.array([1,2,3])`）を引数として、各値に対応する $x^1, x^2, x^3$ の値を並べた NumPy 2次元配列を返す関数 `q21_cubic_features(x)` を作成せよ。

`sklearn.preprocessing.polynomialFeatures` generates features for polynomial regression. Similarly, create a function `q21_cubic_features(x)` that takes a NumPy 1D array (e.g., `np.array([1,2,3])`) as input and returns a 2D NumPy array where each value corresponds to $x^1, x^2, x^3$.

```python
import numpy as np
def q21_cubic_features(x: np.ndarray) -> np.ndarray:
    ...
    return ...

print(q21_cubic_features(np.array([1, 2])))
# The result should be [[1, 1, 1], [2, 4, 8]]
```

```python
grader.check("Exercise 2-1")
```

#### 練習問題 2-2 | 三角関数を用いた基底関数

#### Exercise 2-2 | Base Function with Trigonometric Functions

NumPy 1次元配列（例：`np.array([1,2,3])`）を引数として、各値に対応する $x, \sin(x), \cos(x)$ の値を並べた NumPy 2次元配列を返す関数 `q22_trigonometric_features(x)` を作成せよ。ただし、$\sin(x)$ および $\cos(x)$ はラジアンで計算すること。

Create a function `q22_trigonometric_features(x)` that takes a NumPy 1D array (e.g., `np.array([1,2,3])`) as input and returns a 2D NumPy array where each value corresponds to $x, \sin(x), \cos(x)$. Make sure to compute $\sin(x)$ and $\cos(x)$ in radians.

```python
import numpy as np
def q22_trigonometric_features(x: np.ndarray) -> np.ndarray:
    ...
    return ...

print(q22_trigonometric_features(np.array([0, 1, np.pi])))
# The result should be [[0,0,1], [1, 0.84.., 0.54..], [3.14.., 0, -1]]
```

```python
grader.check("Exercise 2-2")
```

#### 練習問題 2-3 | 三角関数を用いた回帰

#### Exercise 2-3 | Regression Using Trigonometric Functions

練習問題2-2で作成した関数 `q22_trigonometric_features()` を利用して、 $ y = a \sin{x} + b \cos{x} + c x + d $ の回帰を `LinearRegression()` を利用して実施し、その係数を確認せよ。

Using the `q22_trigonometric_features()` function created in Exercise 2-2, perform regression using $ y = a \sin{x} + b \cos{x} + c x + d $ and check the coefficients using `LinearRegression()`.

```python
from sklearn.linear_model import LinearRegression

q23_lr = LinearRegression()
...
print(q23_lr.coef_, q23_lr.intercept_)
```

```python
grader.check("Exercise 2-3")
```

#### 練習問題 2-4 | 二乗平均平方根誤差 (root mean squared error, RMSE) による予測精度評価

#### Exercise 2-4 | Evaluate Prediction Accuracy Using RMSE

練習問題2-3で作成した回帰モデル `q23_lr` について、テストデータ `x_test` および `y_test` に対する二乗平均平方根誤差 (root mean squared error, RMSE) を計算し、 `q24_rmse` に格納せよ。


For the regression model `q23_lr` built in Exercise 2-3, compute the root mean squared error (RMSE) on the test data `x_test` and `y_test`, and store it in `q24_rmse`.

```python
...
q24_rmse = ...
print(q24_rmse)

```

```python
grader.check("Exercise 2-4")
```

Tests for Exercise 1-1 (points: None)
```
>>> isinstance(q11_X, np.ndarray)
True
```
failure_message: q11_X がNumPy配列として定義されていません。 np.array() を使ってNumPy配列に変換してください。 / q11_X is not defined as a NumPy array. Please convert it using np.array().
```
>>> isinstance(q11_y, np.ndarray)
True
```
failure_message: q11_y がNumPy配列として定義されていません。 np.array() を使ってNumPy配列に変換してください。 / q11_y is not defined as a NumPy array. Please convert it using np.array().
```
>>> len(set(q11_y)) == len(q11_y)
True
```
failure_message: q11_y に一致している値が存在します。 / There are duplicate values in q11_y.
```
>>> np.allclose(q11_y, q11_X @ np.array([1, 1]) + 1)
True
```
failure_message: y = x1 + x2 + 1 という等号関係が成り立っていません。 / The equation y = x1 + x2 + 1 is not satisfied.


Tests for Exercise 1-2 (points: None)
```
>>> from sklearn.utils.validation import check_is_fitted
>>> check_is_fitted(q12_lr)
```
failure_message: q12_lr.fit(X,y) が行われていません。学習（回帰モデルの構築）のために fit を行ってください。 / q12_lr.fit(X,y) was not performed. Please perform fit to build the regression model.
```
>>> import numpy as np
>>> __expected_coef = [1.0, 1.0]
>>> np.allclose(q12_lr.coef_, __expected_coef)
True
```
failure_message: q12_lr.coef_ が適切な値になっていません。fitに与えるデータを間違えていませんか？ / q12_lr.coef_ has incorrect values. Are you providing the correct data to fit?
```
>>> import numpy as np
>>> __expected_intercept = 1.0
>>> np.allclose(q12_lr.intercept_, __expected_intercept)
True
```
failure_message: q12_lr.intercept_ が適切な値になっていません。fitに与えるデータを間違えていませんか？ / q12_lr.intercept_ has incorrect values. Are you providing the correct data to fit?


Tests for Exercise 1-3 (points: None)
```
>>> q13_ridge_1.get_params()['alpha'] == 1
True
```
failure_message: q13_ridge_1 において、alpha が適切に設定されていません。Ridgeモデル構築の際に、alpha=1を指定してください。 / alpha is not correctly set for q13_ridge_1. Specify alpha=1 when building the Ridge model.
```
>>> from sklearn.linear_model import Ridge
>>> _ans_q13_ridge_1 = Ridge(alpha=1).fit(q11_X, q11_y)
>>> np.allclose(q13_ridge_1.coef_, _ans_q13_ridge_1.coef_) and np.allclose(q13_ridge_1.intercept_, _ans_q13_ridge_1.intercept_)
True
```
failure_message: q13_ridge_1 において、q13_ridge_1.coef_あるいはq13_ridge_1.intercept_ が誤っています。fitに与えるデータを間違えていませんか？ / q13_ridge_1.coef_ or q13_ridge_1.intercept_ has incorrect values. Are you providing the correct data to fit?
```
>>> q13_ridge_10.get_params()['alpha'] == 10
True
```
failure_message: q13_ridge_10 において、alpha が適切に設定されていません。Ridgeモデル構築の際に、alpha=10を指定してください。 / alpha is not correctly set for q13_ridge_10. Specify alpha=10 when building the Ridge model.
```
>>> from sklearn.linear_model import Ridge
>>> _ans_q13_ridge_10 = Ridge(alpha=10).fit(q11_X, q11_y)
>>> np.allclose(q13_ridge_10.coef_, _ans_q13_ridge_10.coef_) and np.allclose(q13_ridge_10.intercept_, _ans_q13_ridge_10.intercept_)
True
```
failure_message: q13_ridge_10 において、q13_ridge_10.coef_あるいはq13_ridge_10.intercept_ が誤っています。fitに与えるデータを間違えていませんか？ / q13_ridge_10.coef_ or q13_ridge_10.intercept_ has incorrect values. Are you providing the correct data to fit?


Tests for Exercise 2-1 (points: None)
```
>>> isinstance(q21_cubic_features(np.array([1, 2])), np.ndarray)
True
```
failure_message: 返り値が np.array になっていません。 / The return value is not a NumPy array.
```
>>> np.allclose(q21_cubic_features(np.array([1])), np.array([[1, 1, 1]]))
True
```
failure_message: np.array([1]) を与えたときの返り値が誤っています。 / The return value for np.array([1]) is incorrect.
```
>>> np.allclose(q21_cubic_features(np.array([1, 2, 3])), np.array([[1, 1, 1], [2, 4, 8], [3, 9, 27]]))
True
```
failure_message: np.array([1, 2, 3]) を与えたときの返り値が誤っています。x, x^2, x^3 の順番になっていますか？ / The return value for np.array([1, 2, 3]) is incorrect. Is the order x, x^2, x^3?
```
>>> np.allclose(q21_cubic_features(np.array([1, 2.5, 3])), np.array([[1, 1, 1], [2.5, 6.25, 15.625], [3, 9, 27]]))
True
```
failure_message: np.array([1, 2.5, 3]) を与えたときの返り値が誤っています。 / The return value for np.array([1, 2.5, 3]) is incorrect.
```
>>> np.allclose(q21_cubic_features(np.array([-1, 2, -3])), np.array([[-1, 1, -1], [2, 4, 8], [-3, 9, -27]]))
True
```
failure_message: np.array([-1, 2, -3]) を与えたときの返り値が誤っています。 / The return value for np.array([-1, 2, -3]) is incorrect.


Tests for Exercise 2-2 (points: None)
```
>>> isinstance(q22_trigonometric_features(np.array([1, 2])), np.ndarray)
True
```
failure_message: 返り値が np.array になっていません。 / The return value is not a NumPy array.
```
>>> np.allclose(q22_trigonometric_features(np.array([1])), np.array([[1, 0.84, 0.54]]), atol=0.01)
True
```
failure_message: np.array([1]) を与えたときの返り値が誤っています。x, sin x, cos x の順番になっていますか？ / The return value for np.array([1]) is incorrect. Is the order x, sin x, cos x?
```
>>> np.allclose(q22_trigonometric_features(np.array([1, 2, 3])), np.array([[1, 0.84, 0.54], [2, 0.91, -0.42], [3, 0.14, -0.99]]), atol=0.01)
True
```
failure_message: np.array([1, 2, 3]) を与えたときの返り値が誤っています。 / The return value for np.array([1, 2, 3]) is incorrect.
```
>>> np.allclose(q22_trigonometric_features(np.array([1, 2.5, 3])), np.array([[1, 0.84, 0.54], [2.5, 0.6, -0.8], [3, 0.14, -0.99]]), atol=0.01)
True
```
failure_message: np.array([1, 2.5, 3]) を与えたときの返り値が誤っています。 / The return value for np.array([1, 2.5, 3]) is incorrect.
```
>>> np.allclose(q22_trigonometric_features(np.array([-1, 2, -3])), np.array([[-1, -0.84, 0.54], [2, 0.91, -0.42], [-3, -0.14, -0.99]]), atol=0.01)
True
```
failure_message: np.array([-1, 2, -3]) を与えたときの返り値が誤っています。 / The return value for np.array([-1, 2, -3]) is incorrect.


Tests for Exercise 2-3 (points: None)
```
>>> from sklearn.utils.validation import check_is_fitted
>>> check_is_fitted(q23_lr)
```
failure_message: q23_lr.fit(X,y) が行われていません。学習（回帰モデルの構築）のために fit を行ってください。 / q23_lr.fit(X,y) was not performed. Please perform fit to build the regression model.
```
>>> len(q23_lr.coef_) == 3
True
```
failure_message: q23_lr.coef_ が3つの要素からなる配列になっていません。各点はq22_trigonometric_featuresで3つの値で表現されるため、coef_も3つの要素からなる配列になるはずです。 / q23_lr.coef_ does not have 3 elements. Each point is represented by 3 values in q22_trigonometric_features, so coef_ should also be a 3-element array.
```
>>> import numpy as np
>>> __expected_coef = [0.0, 0.0, 1.0]
>>> np.allclose(q23_lr.coef_, __expected_coef)
True
```
failure_message: q23_lr.coef_ が適切な値になっていません。fitに与えるデータを間違えていませんか？ / q23_lr.coef_ has incorrect values. Are you providing the correct data to fit?
```
>>> import numpy as np
>>> __expected_intercept = 0.0
>>> np.allclose(q23_lr.intercept_, __expected_intercept)
True
```
failure_message: q23_lr.intercept_ が適切な値になっていません。fitに与えるデータを間違えていませんか？ / q23_lr.intercept_ has incorrect values. Are you providing the correct data to fit?


Tests for Exercise 2-4 (points: None)
```
>>> __expected_rmse = 0.0
>>> np.allclose(q24_rmse, __expected_rmse)
True
```
failure_message: RMSE値が適切な値になっていません。predictに与えるデータが誤っている、あるいはRMSEの計算方法が誤っている可能性があります。 / The RMSE value is incorrect. You may have provided the wrong data to predict or used an incorrect method for RMSE calculation.
