```python
# Preparation for auto-grading
!pip install otter-grader 
import otter
grader = otter.Notebook()
```

# 基盤データサイエンス演習 第7回 アンサンブル学習 補習資料

# Supplemental Materials: Exercises in Fundamentals on Data Science (7) - Ensemble Learning

## 1 | ブートストラップ・リサンプリング

## 1 | Bootstrap Resampling

#### 事前準備

#### Preliminaries

```python
import numpy as np

q1_X = np.array([
    [ 1,  2,  3],
    [ 4,  5,  6],
    [ 7,  8,  9],
    [10, 11, 12]
])
q1_y = np.array([1, 2, 3, 4])

```

#### 練習問題 1-1 | np.random.choiceによるサンプルの抽出

#### Practice Problem 1-1 | Extracting Samples using np.random.choice

`np.random.choice()`を使って、`q1_X` から復元抽出で2件のサンプルを抽出し、`q11_sampled_X` に記録せよ。ただし、`np.random.seed(42)` を実行してからサンプル抽出を行うこと。

Use `np.random.choice()` to extract 2 samples from `q1_X` with replacement, and record them in `q11_sampled_X`. Be sure to run `np.random.seed(42)` before extracting the samples.

```python
np.random.seed(42)
q11_sampled_X = ...
```

```python
grader.check("Exercise 1-1")
```

#### 練習問題 1-2 | Xとyの対応関係を保持したサンプリングの実施

#### Practice Problem 1-2 | Sampling while Maintaining the Correspondence between X and y

教師あり学習を行う場合、説明変数 `X` と目的変数 `y` の対応関係を保持したままサンプリングを行う必要がある。`np.random.choice()`を使って、`q1_X` と `q1_y` から復元抽出で5件のサンプルを抽出し、`q12_sampled_X` と `q12_sampled_y` に記録せよ。ただし、`np.random.seed(42)` を実行してからサンプル抽出を行うこと。

When performing supervised learning, it is necessary to sample while maintaining the correspondence between the explanatory variable `X` and the target variable `y`. Use `np.random.choice()` to extract 5 samples from `q1_X` and `q1_y` with replacement, and record them in `q12_sampled_X` and `q12_sampled_y`. Be sure to run `np.random.seed(42)` before extracting the samples.

```python
np.random.seed(42)
...
q12_sampled_X = ...
q12_sampled_y = ...

```

```python
grader.check("Exercise 1-2")
```

## 2 | アンサンブル学習における弱学習器 (estimator)

## 2 | Weak Learners (Estimator) in Ensemble Learning

#### 事前準備

#### Preparation


```python
from sklearn import datasets
from sklearn.model_selection import train_test_split
d = datasets.load_digits()
X = d.data
y = d.target
X_train, X_test, y_train, y_test = train_test_split(X, y, train_size=1000, random_state=42)
```

#### 練習問題 2-1 | 精度を最大化する n_estimators の探索：ランダムフォレスト

#### Practice Problem 2-1 | Exploring n_estimators to Maximize Accuracy: Random Forest

`n_estimators` を 1, 2, 4, 10, 20, 40, 100, 200, 400 の9通りに変化させた場合の予測正解率を調査し、全ての予測精度を `q21_all_accuracies` に、最も予測正解率が高かった学習器の数を `q21_best_n_estimators` に、その時の予測正解率を `q21_best_accuracy` に、それぞれ記録せよ。ただし、`random_state=42` を設定すること。

また、付属している描画関数を用いて、`n_estimators` と予測正解率 `q21_all_accuracies` の関係を確認せよ。


Investigate the prediction accuracy when varying `n_estimators` from 1, 2, 4, 10, 20, 40, 100, 200, to 400, and record all prediction accuracies in `q21_all_accuracies`, the number of estimators that achieved the highest accuracy in `q21_best_n_estimators`, and the highest accuracy in `q21_best_accuracy`. Be sure to set `random_state=42`.

Additionally, use the provided plotting function to visualize the relationship between `n_estimators` and prediction accuracy `q21_all_accuracies`.

```python
from sklearn.ensemble import RandomForestClassifier
from sklearn.metrics import accuracy_score
q21_all_accuracies = []
q21_best_n_estimators = None
q21_best_accuracy = 0.0

for n_estimators in [1,2,4,10,20,40,100,200,400]:
    ...
```

```python
import matplotlib.pyplot as plt
plt.plot([1,2,4,10,20,40,100,200,400], q21_all_accuracies)
plt.xscale("log")
plt.grid()
plt.show()
```

```python
grader.check("Exercise 2-1")
```

Tests for Exercise 1-1 (points: None)
```
>>> q11_sampled_X.shape[0] == 2
True
```
failure_message: q11_sampled_X のサンプル数が正しくありません。
```
>>> q11_sampled_X.shape[0] == 2
True
```
failure_message: The number of samples in q11_sampled_X is incorrect.
```
>>> np.all(q11_sampled_X == np.array([2, 3]))
np.True_
```
failure_message: q11_sampled_X の値が誤っています。復元抽出をおこなっていますか？np.random.seed(42)は最初に実行されていますか？
```
>>> np.all(q11_sampled_X == np.array([2, 3]))
np.True_
```
failure_message: The values of q11_sampled_X are incorrect. Are you performing resampling? Has np.random.seed(42) been run first?


Tests for Exercise 1-2 (points: None)
```
>>> q12_sampled_X.shape[0] == 5 and q12_sampled_X.shape[1] == 3
True
```
failure_message: q12_sampled_X のサンプル数が正しくありません。
```
>>> q12_sampled_X.shape[0] == 5 and q12_sampled_X.shape[1] == 3
True
```
failure_message: The number of samples in q12_sampled_X is incorrect.
```
>>> q12_sampled_y.shape[0] == 5
True
```
failure_message: q12_sampled_y のサンプル数が正しくありません。
```
>>> q12_sampled_y.shape[0] == 5
True
```
failure_message: The number of samples in q12_sampled_y is incorrect.
```
>>> boolean = True
>>> for col in range(q12_sampled_X.shape[0]):
...     __match_index = np.all(q1_X == q12_sampled_X[col], axis=1)
...     boolean = boolean == __match_index.any()
...     __ind = np.where(__match_index)
...     __ytest = (q1_y[__ind] == q12_sampled_y[col]).any()
...     boolean = boolean == __ytest
>>> boolean == True
np.True_
```
failure_message: q12_sampled_X と q12_sampled_y の対応関係が保たれていません。
```
>>> boolean = True
>>> for col in range(q12_sampled_X.shape[0]):
...     __match_index = np.all(q1_X == q12_sampled_X[col], axis=1)
...     boolean = boolean == __match_index.any()
...     __ind = np.where(__match_index)
...     __ytest = (q1_y[__ind] == q12_sampled_y[col]).any()
...     boolean = boolean == __ytest
>>> boolean == True
np.True_
```
failure_message: The correspondence between q12_sampled_X and q12_sampled_y is not maintained.
```
>>> np.all(q12_sampled_X == np.array([[7, 8, 9], [10, 11, 12], [1, 2, 3], [7, 8, 9], [7, 8, 9]])) and np.all(q12_sampled_y == np.array([3, 4, 1, 3, 3]))
np.True_
```
failure_message: q12_sampled_X と q12_sampled_y の値が誤っています。復元抽出をおこなっていますか？np.random.seed(42)は最初に実行されていますか？
```
>>> np.all(q12_sampled_X == np.array([[7, 8, 9], [10, 11, 12], [1, 2, 3], [7, 8, 9], [7, 8, 9]])) and np.all(q12_sampled_y == np.array([3, 4, 1, 3, 3]))
np.True_
```
failure_message: The values of q12_sampled_X and q12_sampled_y are incorrect. Are you performing resampling? Has np.random.seed(42) been run first?


Tests for Exercise 2-1 (points: None)
```
>>> q21_best_n_estimators == 100
True
```
failure_message: q21_best_n_estimators の値が誤っています。random_state=42 で fit していますか？ / The value of q21_best_n_estimators is incorrect. Are you fitting with random_state=42?
```
>>> np.isclose(q21_best_accuracy, 0.973651191969887, atol=1e-05)
np.True_
```
failure_message: q21_best_accuracy の値が誤っています。random_state=42 で fit していますか？ / The value of q21_best_accuracy is incorrect. Are you fitting with random_state=42?
```
>>> len(q21_all_accuracies) == 9
True
```
failure_message: 試行回数が足りていません。[1,2,4,10,20,40,100,200,400]の全ての要素に対して試行していますか？ / Not enough trials were performed. Have you tried all elements of [1,2,4,10,20,40,100,200,400]?
```
>>> np.allclose(q21_all_accuracies, [0.7728983688833124, 0.7540777917189461, 0.8782936010037641, 0.946047678795483, 0.9598494353826851, 0.9711417816813049, 0.973651191969887, 0.9698870765370138, 0.972396486825596], atol=1e-05)
True
```
failure_message: q21_all_accuracies の値が誤っています。random_state=42 で fit していますか？ / The values of q21_all_accuracies are incorrect. Are you fitting with random_state=42?
